# Programming in Python

__Laboratory 02, Type system and built-in types__

__Sept 27, 2022__

__Submission deadline: Sept 29 23:59__

__Submissions must be made through Github Classroom__

Some exercises are graded.
The maximum score you can get is 25 points.
The score value of graded exercises is listed in the instructions.

You are free to use any material (lecture, Stackoverflow etc.) except full solutions but you should first try to figure them out on your own.

We use an [nbgrader](https://nbgrader.readthedocs.io/en/stable/index.html) for autograding the assignments. Many cells are read-only. Please **DO NOT** delete or copy cells. You can add new cells but your solutions should go into the designated cells.

We will use the wikipedia module for some exercises. It can be installed by running pip from Jupyter:

The Python Standard Library has solutions for many of these tasks but you should solve them without relying on the standard library.

We will use the `wikipedia` module in notebook.

```python
! pip install wikipedia
```

Output:
```
Collecting wikipedia
  Downloading wikipedia-1.4.0.tar.gz (27 kB)
  Preparing metadata (setup.py): started
  Preparing metadata (setup.py): finished with status 'done'
Collecting beautifulsoup4
  Downloading beautifulsoup4-4.12.2-py3-none-any.whl (142 kB)
Requirement already satisfied: requests<3.0.0,>=2.0.0 in c:\users\user\appdata\roaming\python\python38\site-packages (from wikipedia) (2.23.0)
Requirement already satisfied: urllib3!=1.25.0,!=1.25.1,<1.26,>=1.21.1 in c:\users\user\appdata\roaming\python\python38\site-packages (from requests<3.0.0,>=2.0.0->wikipedia) (1.25.9)
Requirement already satisfied: chardet<4,>=3.0.2 in c:\users\user\appdata\roaming\python\python38\site-packages (from requests<3.0.0,>=2.0.0->wikipedia) (3.0.4)
Requirement already satisfied: idna<3,>=2.5 in c:\users\user\appdata\roaming\python\python38\site-packages (from requests<3.0.0,>=2.0.0->wikipedia) (2.9)
Requirement already satisfied: certifi>=2017.4.17 in c:\users\user\appdata\roaming\python\python38\sit
```

# 1. Define a function that takes a sequence as its input and returns whether the sequence is symmetric. A sequence is symmetric if it is equal to its reverse.

<!-- grade_id: cell-9063e2fba9c7157d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def is_symmetric(sequence):
    # Check if the sequence is equal to its reverse
    return sequence == sequence[::-1]


```

<!-- grade_id: cell-58b3ad7e361196e7, points: 0, grade: true, solution: false -->
```python
assert is_symmetric([1]) is True
assert is_symmetric([None, None]) is True
assert is_symmetric([]) is True
assert is_symmetric([1, 2, 3, 1]) is False
assert is_symmetric([1, "foo", "bar", "foo", 1]) is True
assert is_symmetric("abcba") is True
assert is_symmetric(list(range(100)) + list(range(99, -1, -1))) is True
assert is_symmetric(list(range(100)) + list(range(100, 0, -1))) is False
assert is_symmetric([1, "1"]) is False
```

# String manipulation

## 2.1 Define a function that takes a string as its input and return a dictionary with the character frequencies.

Do not use the `collections` module.

<!-- grade_id: cell-7fbaf8610be225a9, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def char_freq(input_string):
    char_count = {}

    for char in input_string:
        if char in char_count:
            char_count[char] += 1
        else:
            char_count[char] = 1

    return char_count

```

<!-- grade_id: cell-ecedb93a6c391539, points: 1, grade: true, solution: false -->
```python
assert char_freq("aba") == {"a": 2, "b": 1}
assert char_freq("") == {}
assert char_freq("ab" * 100 + "cd") == {"a": 100, "b": 100, "c": 1, "d": 1}
```

## 2.2 (1p) Define a function that computes word frequencies in a text.

You can split the text with `text.split()`.

<!-- grade_id: cell-4d4d473dd807f962, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def word_freq(text):
    word_count = {}
    words = text.split()

    # Iterate through the words
    for word in words:
        cleaned_word = word.strip('.,!?"\'').lower()

        if cleaned_word in word_count:
            word_count[cleaned_word] += 1
        else:
            word_count[cleaned_word] = 1

    return word_count

```

<!-- grade_id: cell-1c299c2c935129e7, points: 0, grade: true, solution: false -->
```python
s = "the green tea and the black tea"
assert word_freq(s) == {"the": 2, "tea": 2, "green": 1, "black": 1, "and": 1}
assert word_freq("") == {}
```

## 2.3 (2p) Define a function that takes two strings and decides whether they are anagrams. A string is an anagram of another string if its letters can be rearranged so that it equals the other string.¶

For example:

```
abc -- bac
aabb -- abab
```

Counter examples:

```
abc -- aabc
abab -- aaab
```

<!-- grade_id: cell-7ea849b6ffc4421a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def are_anagrams(str1, str2):
    str1 = str1.replace(" ", "").lower()
    str2 = str2.replace(" ", "").lower()

    if len(str1) != len(str2):
        return False

    char_count1 = {}
    char_count2 = {}

    for char in str1:
        if char in char_count1:
            char_count1[char] += 1
        else:
            char_count1[char] = 1

    for char in str2:
        if char in char_count2:
            char_count2[char] += 1
        else:
            char_count2[char] = 1

    return char_count1 == char_count2
```

<!-- grade_id: cell-231f78d971a5d48b, points: 2, grade: true, solution: false -->
```python
assert are_anagrams("", "") is True
assert are_anagrams("abc", "bac") is True
assert are_anagrams("aabb", "abab") is True
assert are_anagrams("abab", "aaab") is False
```

## 2.4 (2p) Define a function that takes a string and returns its bigram frequencies as a dictionary.

Character bigrams are pairs of subsequent characters. For example word apple contains the following bigrams: ap, pp, pl, le.

Do not use the `collections` module.

<!-- grade_id: cell-2669edc3d29b28b6, points: 0, grade: false, solution: true -->
```python
def get_char_bigrams(text):
    bigram_freq = {}
    
    for i in range(len(text) - 1):
        bigram = text[i:i + 2] 
        if bigram in bigram_freq:
            bigram_freq[bigram] += 1
        else:
            bigram_freq[bigram] = 1
    
    return bigram_freq

```

<!-- grade_id: cell-5a3764a1d0909393, points: 2, grade: true, solution: false -->
```python
assert get_char_bigrams("apple") == {"ap": 1, "pp": 1, "pl": 1, "le": 1}
assert get_char_bigrams("apple apple") == {
    "ap": 2,
    "pp": 2,
    "pl": 2,
    "le": 2,
    "e ": 1,
    " a": 1,
}
```

# Wikipedia module

The following exercises use the `wikipedia` package. The basic usage is illustrated below.

The documentation is available [here](https://pypi.python.org/pypi/wikipedia/).

Searching for pages:

```python
import wikipedia

results = wikipedia.search("Vienna")
results
```

Output:
```
['Vienna',
 'Battle of Vienna',
 'Vienna Convention',
 'Vienna Award',
 'Vienna Teng',
 'Vienna sausage',
 'Siege of Vienna',
 'Archbishop of Vienna',
 'Congress of Vienna',
 'Treaty of Vienna']
```

Downloading an article:

```python
article = wikipedia.page("Budapest")

article.summary[:200]
```

Output:
```
'Budapest (UK: , US: ; Hungarian pronunciation: [ˈbudɒpɛʃt] ) is the capital and most populous city of Hungary. It is the ninth-largest city in the European Union by population within city limits and t'
```

The content attribute contains the full text:

```python
type(article.content), len(article.content)
```

Output:
```
(str, 86879)
```

By default the module downloads the English Wikipedia. The language can be changed the following way:

```python
wikipedia.set_lang("de")
```

```python
wikipedia.search("Budapest")
```

Output:
```
['Budapest',
 'Budapester Memorandum',
 'Grand Budapest Hotel',
 'Schlacht um Budapest',
 'Parlamentsgebäude (Budapest)',
 'Budapester',
 'Bahnhof Budapest-Keleti',
 'Kettenbrücke (Budapest)',
 'Ungarn',
 'Metró Budapest']
```

```python
de_article = wikipedia.page("Budapest")
de_article.summary[:100]
```

Output:
```
"Budapest (ungarische Aussprache ['budɒpɛʃt]; ; deutsch historisch Ofen-Pesth) ist die Hauptstadt und"
```

```python
wikipedia.set_lang("en")
```

## 3.1 (5p) Build a small Wikipedia corpus by downloading pages until their character count exceeds 100000. The corpus should be a list of strings.

You can download arbitrary pages but some pages redirect to disambiguation pages. You should not pick these. Make sure the same page is not included multiple times.

<!-- grade_id: cell-2e7be91b99158dca, points: 0, grade: false, solution: true -->
```python
import wikipedia

def build_wikipedia_corpus(max_chars=100000):
    wikipedia.set_lang("en")
    
    corpus = []
    visited_pages = set()
    total_chars = 0

    while total_chars < max_chars:
        page_title = wikipedia.random()
        if page_title in visited_pages:
            continue

        try:
            page = wikipedia.page(page_title)
            content = page.content
            if content:
                corpus.append(content)
                visited_pages.add(page_title)
                total_chars += len(content)
        except wikipedia.exceptions.PageError:
            continue
        except wikipedia.exceptions.DisambiguationError:
            continue

    return corpus



```

<!-- grade_id: cell-f26aa917928c62c8, points: 5, grade: true, solution: false -->
```python
corpus = build_wikipedia_corpus()
# check types
assert isinstance(corpus, list)
assert isinstance(corpus[0], str)

# check character count
sum_length = 0
for document in corpus:
    sum_length += len(document)
assert sum_length >= 100000

# no duplicate pages
assert len(set(corpus)) == len(corpus)
```

## 3.2 (2p) Add a `min_char` parameter to `build_wikipedia_corpus` that defaults to 10000. You should stop building the corpus when this number is reached.

<!-- grade_id: cell-bf7ffeba0644e753, points: 0, grade: false, solution: true -->
```python
import wikipedia

def build_wikipedia_corpus(max_chars=100000, min_chars=10000, max_retries=1):
    wikipedia.set_lang("en")
    
    corpus = []
    visited_pages = set()
    total_chars = 0

    while total_chars < max_chars:
        retries = 0
        while retries < max_retries:
            try:
                page_title = wikipedia.random()
                if page_title in visited_pages:
                    continue

                page = wikipedia.page(page_title)
                content = page.content

                if content and len(content) >= min_chars:
                    if total_chars + len(content) <= max_chars:
                        corpus.append(content)
                        visited_pages.add(page_title)
                        total_chars += len(content)
                        break  
                    else:
                        break  
            except (wikipedia.exceptions.DisambiguationError, wikipedia.exceptions.PageError):
                retries += 1

    return corpus



```

```python
corpus = build_wikipedia_corpus(1000)
# check types
assert isinstance(corpus, list)
assert isinstance(corpus[0], str)

# check character count
sum_length = 0
for document in corpus:
    sum_length += len(document)
assert sum_length >= 1000
# shouldn't be too long
assert sum_length < 100000
```

## 3.3 (3p) Compute the word frequencies of your Wikipedia corpus. Return the $k$ most frequent words with their frequencies as a list of tuples. $k$ should default to 10. The list should be ordered by frequency starting from the most frequent.

You can use str.split() for word tokenization.

<!-- grade_id: cell-8c167ceb2a4c65f1, points: 0, grade: false, solution: true -->
```python
from collections import Counter

def get_corpus_word_frequencies(corpus, k=10):
    word_counter = Counter()
    for text in corpus:
        words = text.split()  
        word_counter.update(words)  
    
    most_common_words = word_counter.most_common(k)
    
    return most_common_words

```

<!-- grade_id: cell-d17a9e499c13bd35, points: 3, grade: true, solution: false -->
```python
top_words_and_freqs = get_corpus_word_frequencies(corpus)
assert isinstance(top_words_and_freqs, list)
assert len(top_words_and_freqs) == 10
assert isinstance(top_words_and_freqs[0], tuple)

fr = top_words_and_freqs[0][1]

for word, freq in top_words_and_freqs:
    assert isinstance(word, str)
    assert isinstance(freq, int)
    assert freq <= fr
    fr = freq

top_words_and_freqs5 = get_corpus_word_frequencies(corpus, k=5)
assert len(top_words_and_freqs5) == 5
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-77-25f165b5151a> in <module>
      1 top_words_and_freqs = get_corpus_word_frequencies(corpus)
      2 assert isinstance(top_words_and_freqs, list)
----> 3 assert len(top_words_and_freqs) == 10
      4 assert isinstance(top_words_and_freqs[0], tuple)
      5 

AssertionError:
```

## 3.4.2 Print the ten most frequent words and their frequency in the following format:

```
a<TAB>12
an<TAB>11
```

<!-- grade_id: cell-bc2d984be320ff7c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

corpus_word_frequencies = get_corpus_word_frequencies(corpus, k=10)

for word, frequency in corpus_word_frequencies:
    print(f"{word}\t{frequency}")

```

# Linear algebra

## 4.1 (2p) Define a function that takes a list of lists and decided if it is a matrix. A list of lists is a matrix if each nested list is the same length.

<!-- grade_id: cell-90b341f4b426d33d, points: 0, grade: false, solution: true -->
```python
def is_matrix(matrix):
    # Check if the matrix is empty
    if not matrix:
        return False

    first_row_length = len(matrix[0])

    for row in matrix:
        if len(row) != first_row_length:
            return False

    return True


```

<!-- grade_id: cell-218d090c410748b5, points: 2, grade: true, solution: false -->
```python
#assert is_matrix([]) is True
assert is_matrix([[], []]) is True
assert is_matrix([[1, 2], [2, 1]]) is True
assert is_matrix([[1], [2, 1]]) is False
assert is_matrix([[1, 2], [2, 1], [3, 4]]) is True
assert is_matrix([[1, 2], [2, 1], [3, 4, 5]]) is False
```

## 4.2 (2p) Define a function that takes a matrix as an input represented as a list of lists. Return its transpose without changing the original matrix.

Raise a `ValueError` if the input is not a valid matrix. Don't forget to reuse your code.

<!-- grade_id: cell-adee24575dcbac8e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def is_valid_matrix(matrix):
    if not matrix or not all(isinstance(row, list) for row in matrix):
        return False
    num_columns = len(matrix[0])
    return all(len(row) == num_columns for row in matrix)

def transpose(matrix):
    if not is_valid_matrix(matrix):
        raise ValueError("Input is not a valid matrix")
    
    num_rows = len(matrix)
    num_columns = len(matrix[0])
    transposed = [[matrix[j][i] for j in range(num_rows)] for i in range(num_columns)]
    
    return transposed

```

<!-- grade_id: cell-9b412cb3d0c0506a, points: 2, grade: true, solution: false -->
```python
m1 = [[1, 2, 3], [4, 5, 6]]
m2 = [[1, 4], [2, 5], [3, 6]]

assert transpose([[1]]) == [[1]]
assert transpose(m1) == m2
# m1 is unchanged
assert m1 == [[1, 2, 3], [4, 5, 6]]

# transpose is an involution (self-inverse)
assert transpose(transpose(m1)) == m1
```

## 4.3 (1p) Define a function that takes a matrix as its input and returns the row-wise maximum elements as a list.

Raise a `ValueError` if the input is not a valid matrix. Don't forget to reuse your code.

<!-- grade_id: cell-21f91aada57c0332, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def rowwise_max(matrix):
    if not is_valid_matrix(matrix):
        raise ValueError("Input is not a valid matrix")
    
    max_elements = []
    for row in matrix:
        max_value = max(row)
        max_elements.append(max_value)
    
    return max_elements

```

<!-- grade_id: cell-eb440f91ea977032, points: 1, grade: true, solution: false -->
```python
assert rowwise_max(
    [
        [1, 2, 3],
        [1, 4, 3],
        [1, 5, 3],
    ]
) == [3, 4, 5]

assert rowwise_max([list(range(1000))]) == [999]
assert rowwise_max(transpose([list(range(1000))])) == list(range(1000))
```

## 4.4 (2p) Define a function that takes a matrix as its input and returns the column-wise maximum elements as a list.

Try to reuse your previous functions.

<!-- grade_id: cell-bde0cb9cf8d75c89, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def columnwise_max(matrix):
    if not is_valid_matrix(matrix):
        raise ValueError("Input is not a valid matrix")
    
    # Transpose the matrix to treat columns as rows
    transposed_matrix = transpose(matrix)
    
    max_elements = []
    for row in transposed_matrix:
        max_value = max(row)
        max_elements.append(max_value)
    
    return max_elements

```

<!-- grade_id: cell-142e6034ed15005b, points: 2, grade: true, solution: false -->
```python
assert columnwise_max(
    [
        [1, 2, 3],
        [1, 4, 3],
        [1, 5, 3],
    ]
) == [1, 5, 3]

assert columnwise_max([list(range(1000))]) == list(range(1000))
assert columnwise_max(transpose([list(range(1000))])) == [999]
```

## 4.5 (3p) Define a function that takes two vectors represented as lists and returns their L2 distance.

Also called Euclidean distance.

<!-- grade_id: cell-849b6ea4c65f058e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
import math

def l2_distance(vector1, vector2):
    if len(vector1) != len(vector2):
        raise ValueError("Input vectors must have the same length")
    
    sum_of_squares = sum((x - y) ** 2 for x, y in zip(vector1, vector2))
    distance = math.sqrt(sum_of_squares)
    
    return distance

```

<!-- grade_id: cell-4c9f06a8534cbefc, points: 3, grade: true, solution: false -->
```python
assert abs(l2_distance([1, 0], [1, 0]) - 0) < 1e-5
assert abs(l2_distance([1, 0], [0, 1]) - 2**0.5) < 1e-5
```

## 4.6 Create a function that takes a matrix of size NxM and returns the pairwise L2 distance of each row pair.

For input $A^{N \times M}$, the output should be $D^{N \times N}$ where: $D_{i,j} = L_2 (A_i, A_j)$

<!-- grade_id: cell-644d3391fc995ee9, points: 0, grade: false, solution: true -->
```python
def pairwise_row_distance(matrix):
    if not is_valid_matrix(matrix):
        raise ValueError("Input is not a valid matrix")
    
    num_rows = len(matrix)
    distances = [[0.0] * num_rows for _ in range(num_rows)]
    
    for i in range(num_rows):
        for j in range(i, num_rows):
            distance = l2_distance(matrix[i], matrix[j])
            distances[i][j] = distance
            distances[j][i] = distance  # Since L2 distance is symmetric
    
    return distances

```

<!-- grade_id: cell-8bc3cf03b303c7c3, points: 0, grade: true, solution: false -->
```python
D = pairwise_row_distance([[1, 0, 2], [0, 2, 1]])
# should be a matrix
assert is_matrix(D) is True
assert len(D) == 2
assert len(D[0]) == 2
assert D[0][0] == 0
assert abs(D[0][1] - 2.449489742783178) < 1e-3
assert abs(D[1][0] - 2.449489742783178) < 1e-3
assert D[1][1] == 0

# should be symmetric
assert transpose(D) == D
```

```python

```